---
title: 'Working with the RStudio IDE (Part 1)'
description: 'In this chapter, you will learn to get familiar with R and RStudio.  Garrett will take you through steps of downloading R to being familiar with the RStudio environment.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
